<div class = "alert alert-warning" align="center"> 
    <img src="https://www.egd.mg/wp-content/uploads/2020/01/cropped-logo-home-110px.png" align="left" width="290"  margin="50px"> 
    <center><h1>Lycée français de Tananarive</h1> 
        <h2> Jupyter Hub,  Spécialité NSI </h2></center>
</div>

# Utiliser les listes
## 1. Position du problème
Lorsque l'on réunit **23** personnes, il y a une chance sur deux pour que deux d'entre elles soient nées le même jour. Ce résultat contre intuitif est appelé le *paradoxe des anniversaires*.

Nous allons écrire une suite de scripts en Python qui le vérifie expérimentalement.

## 2. Vérification expérimentale du paradoxe des anniversaires.
## 2.1 Une première solution
### 2.1.1 Génération de la liste des anniversaires
Générer un tableau de taille 23 avec des dates d'anniversaires tirées au hasard. les dates seront un nombre compris entre 1 et 365. Commençons par importer la bibliothéque random : 

```python
from random import randint
```

Et l'utiliser pour écrire une fonction `genere_dates()` qui renvoie une liste de 23 entiers compris entre 1 et 365.

<!-- grade_id: cell-28234ff291c601e2, points: 0, grade: false, solution: true -->
```python
def genere_dates() :
    # YOUR CODE HERE
    return [randint(1,365) for i in range(23)]
```

<!-- grade_id: cell-4e16daebe01dd1fc, points: 2, grade: true, solution: false -->
```python
assert len(genere_dates()) == 23
assert max(genere_dates()) <= 365
assert min(genere_dates()) >= 1
```

### 2.1.2 Vérification de la présence d'un doublet.
On désire maintenant écrire une fonction `verification(tab)`qui vérifie si il y a deux dates d'anniversaire identiques dans une liste passée en argument à la fonction. Pour cela on comparera si le premier nombre est identique au suivant, puis on comparera le second élément aux suivant etc ...

**Spécification de `verification(tab)`**
- Données :
    * tab : Une liste de taille variable de nombres entiers.
- Valeur retournée :
    * True : si la liste contient deux fois le même élèment
    * False : si la liste ne contient pas deux fois le même élément.


<!-- grade_id: cell-421f48ab5b5bad13, points: 0, grade: false, solution: true -->
```python
def verification(tab):
    # YOUR CODE HERE
    for i in range(len(tab)):
        if (nb_occurences(tab[i],tab) > 1):
            return True
    return False
def nb_occurences(n,tab) :
    nbre = 0
    # YOUR CODE HERE
    for i in tab:
        if (n == i):
            nbre = nbre + 1
    return nbre
```

```python
# ou encore
def verif(tab):
    for x in tab:
        if set([num for num in tab if tab.count(x) > 1]) != set():
            return True
        return False
```

<!-- grade_id: cell-b4021bfbf6ff1810, points: 2, grade: true, solution: false -->
```python
assert verification([1,2,1,2]) == True
assert verification([1,2,3,4,5,6,7,8,9]) == False
assert verification([1,2,2,2,2,2,2,1]) == True
assert verification([5,4,3,2,1,0]) == False
```

### 2.1.3 Vérifications expérimentale du paradoxe des anniversaires.
Exécuter la fonction `paradox(n)` qui vérifie n fois si une liste de 23 dates contient un doublons. La fonction retourne le nombre de fois ou un doublon est rencontré.

```python
def paradox(n) :
    cpt = 0
    for i in range(n) :
        tab = genere_dates()
        if verification(tab) == True :
            cpt += 1
    return(cpt)
```

```python
print (" Pour 500 listes, on trouve ",paradox(500)," listes qui contiennent un doublon")
print (" Pour 1000 listes, on trouve ",paradox(1000)," listes qui contiennent un doublon")
print (" Pour 10000 listes, on trouve ",paradox(10000)," listes qui contiennent un doublon")
```

Output:
```
Pour 500 listes, on trouve  250  listes qui contiennent un doublon
 Pour 1000 listes, on trouve  519  listes qui contiennent un doublon
 Pour 10000 listes, on trouve  5170  listes qui contiennent un doublon
```

```python

```